<a href="https://colab.research.google.com/github/yaminichowdaryu/CE880/blob/main/week10/problem10.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-10 : Problem notebook

`10-fold cross-validation` is a common technique used in machine learning to assess the performance of a predictive model. In this method, the dataset is randomly partitioned into 10 equal-sized folds. The model is trained and evaluated 10 times, each time using a different fold as the validation set and the remaining 9 folds as the training set. This process ensures that each data point is used for validation exactly once.

---
## Question 1:

Write a function to perform 10-fold cross-validation to evaluate the performance of a predictive model and return meam accuracy. You are free to use the following two function:

 * `from sklearn.model_selection import KFold`
 * `from sklearn.metrics import accuracy_score`

 In the test cells, we have loaded two datasets for you: 1) iris data `load_iris` and training a `RandomForestClassifier`; and 2) Breast cancer data `load_breast_cancer` and training a `SVC`.
 So, create function `cross_validation(model, X, y, n_splits=10)` keep test cells mind.

<!-- grade_id: my_model_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import KFold
from sklearn.metrics import accuracy_score

def cross_validation(model, X, y, n_splits=10):
    """
    Perform 10-fold cross-validation to evaluate the performance of a predictive model.

    Parameters:
    - model: Machine learning model to be evaluated.
    - X (numpy.ndarray): Feature data.
    - y (numpy.ndarray): Target labels.
    - n_splits (int): Number of splits for cross-validation.

    Returns:
    - mean_accuracy (float): Mean accuracy across all folds.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: my_model_test_RF, points: 4, grade: true, solution: false -->
```python
# Check you solution by running this cell Test case 1: this cell of 4 marks
import math
from sklearn.datasets import load_iris
from sklearn.ensemble import RandomForestClassifier

# Load Iris dataset
iris = load_iris()
X = iris.data
y = iris.target

# Initialize a Random Forest classifier
model = RandomForestClassifier(n_estimators=100, random_state=42)

# Perform 10-fold cross-validation
mean_accuracy = cross_validation(model, X, y)

print("Mean Accuracy:", mean_accuracy)

assert math.isclose(mean_accuracy, 0.96, rel_tol = 0.01)
```

<!-- grade_id: my_model_test_SVM, points: 5, grade: true, solution: false -->
```python
# Check you solution by running this cell Test case 1: this cell of 5 marks
from sklearn.datasets import load_breast_cancer
from sklearn.svm import SVC
from sklearn.preprocessing import StandardScaler

# Load Breast Cancer dataset
breast_cancer = load_breast_cancer()
X = breast_cancer.data
y = breast_cancer.target

# Scale the features
scaler = StandardScaler()
X_scaled = scaler.fit_transform(X)

# Initialize a Support Vector Classifier
model = SVC(kernel='linear', C=1.0, random_state=42)

# Perform 10-fold cross-validation
mean_accuracy = cross_validation(model, X_scaled, y)
print("Mean Accuracy:", mean_accuracy)

assert math.isclose(mean_accuracy, 0.97, rel_tol = 0.01)
```